# Exercise Sheet 4
## Instructions
The instructions for this exercise sheet are the same as the previous, so this section is just a summary to serve as a reminder. Refer to the material on Engage and at the start of Exercise Sheet 1 for more detail, and if you are unsure ask on the Q&A forum.

This exercise sheet counts towards your overall grade for the unit.

Complete each question in the code cell provided, which will usually contain some skeleton code to help you start the question. Unless specified otherwise you may change the skeleton code as you wish, but your code must pass the formatting tests in the following code cell – if the cell runs without errors then your code is eligible for submission. However these will only test the basics of the question. Your code will be subject to additional hidden tests which will determine your grade. Make sure you thoroughly test your own code to check you think it works correctly before submission.

## Questions
### Question 1 & 2
*(worth 25% each)*

In the cell below, implement 
1. the *bubble sort* and 
2. the *merge sort* algorithms.

Unlike the examples in the unit material, please make your functions *return new sorted versions* of the input lists, *without* modifying the input lists. In other words, do not worry about making your algorithms in-place.

If you find it easier to think about modifying the list in-place (more likely with bubble sort than merge sort), you can always write `new_list = in_list.copy()` on the first line, sort `new_list` in place, then return `new_list`.

#### Bubble Sort
Bubble sort works by iterating over the list, comparing adjacent items, and swapping them if they are out of order. 

Notice that after the first pass of a bubble sort, the biggest item is always moved to the end of the list. This means the next iteration does not need to check the final position of the list. 

In addition, you can keep track of how many items were swapped on each iteration. If this count is zero on any iteration, then the list must be fully sorted, in which case the algorithm can stop early.

#### Merge Sort
To understand merge sort, let's first understand the *merge* operation. The merge operation takes two *sorted* lists and combines them into another sorted list. Imagine we are merging `list1` and `list2`, both of which are sorted. Create two variables, one pointing to the first element (zero) of each list, call this `ptr1` and `ptr2`.

Then, go through each index of the merged list (e.g. `i in range(0, len(list1)+len(list2))`). Set position `i` of the merged list to the *smaller* value of `list1[ptr1]` and `list2[ptr2]`, then advance the corresponding pointer by one.

So initially you will compare the first element of each list, if the value in `list1` is smaller then copy this into the merged list, and then compare the second item of `list1` with the first item of `list2`.

Keep going until one pointer goes beyond its list – at this point, copy the rest of the other list into the merged list.

Have a look at an example of the merge operation below:

<br /><video controls loop autoplay width=600 src="./resources/merge.mp4">
</video>

Now, to understand *merge sort*, realise that any list of size 1 must be considered sorted. So, here is the process for the input list:
* If `in_list` is length 1, return `in_list`
* Otherwise, split `in_list` list into two halves: `left` and `right`
* Merge sort `left`
* Merge sort `right`
* Merge `left` and `right` into `merged`
* Return `merged`

For any size input list, this will be recursively split into two until both halves are length 1, and then the final sorted list will be rebuilt by merging sorted lists.

<!-- grade_id: cell-ffd4978cfbf4832e, points: 0, grade: false, solution: true -->
```python
def bubble_sort(in_list):
    new_list = in_list.copy()
    total = len(new_list)
    for i in range(total):
        swap = False
        for j in range(0, total-i-1):           # -i sine the last i elements are already sorted
            if new_list[j] > new_list[j+1]:
                new_list[j], new_list[j+1] = new_list[j+1], new_list[j]
                swap = True
        if not swap:
            break                               # if no swaps on this run then the list is already sorted and we can stop 
    return new_list
    
def merge_sort(in_list):
    if len(in_list) <=1:
        return in_list                          # Base case where a list 1 elements is already sorted and list of 0 elements doesn't need sorting
    
    middle = len(in_list) // 2
    left = merge_sort(in_list[:middle])         # Recursively sorting the left half of the list
    right = merge_sort(in_list[middle:])        # Recursively sorting the right half of the list
    return merge(left, right)

def merge(left, right):
    merged_list = []
    i = j = 0
    while i < len(left) and j < len(right):
        if left[i] <= right[j]:
            merged_list.append(left[i])
            i += 1
        else:
            merged_list.append(right[j])
            j += 1
    # Append any of the remaining items from either list
    merged_list += left[i:]
    merged_list += right[j:]
    return merged_list
```

<!-- grade_id: cell-e322586715be1ec0, points: 50, grade: true, solution: false -->
```python
# Note that your code will be checked to ensure you are implementing the correct algorithms!
assert(bubble_sort([37, 42, 9, 19, 35, 4, 53, 22]) == [4, 9, 19, 22, 35, 37, 42, 53])
assert(bubble_sort([5, 4, 3, 2, 1]) == [1, 2, 3, 4, 5])
assert(bubble_sort([]) == [])

assert(merge_sort([37, 42, 9, 19, 35, 4, 53, 22]) == [4, 9, 19, 22, 35, 37, 42, 53])
assert(merge_sort([5, 4, 3, 2, 1]) == [1, 2, 3, 4, 5])
assert(merge_sort([]) == [])
```

### Question 3
*(worth 25%)*

In the markdown cell below, in your own words, explain the time complexity, space complexity, and stability of bubble sort and merge sort – both as algorithms, and your specific implementations.

Of course, these are well known algorithms, so we are aware you can look up the answers. The marks in this section are awarded for for explaining *why* these properties hold, in your own words. I encourage you to try to work them out yourself before you look up the answers, this might make explaining it easier.

<!-- grade_id: cell-c50e99cdef82c2d1, points: 25, grade: true, solution: true -->
YOUR ANSWER HERE:
### Bubble Sort:

#### Time complexity
Time complexity for the **worst case** is `O(n^2)` as you can't just loop over every element of the list once as that would just push the largest item to the right, you need to go through several iterations to move each larger item to the right; this could be up to `total` moves where total is `n` so because in the worst case inner loop and outer loop would review each item up `n items` and have to swap on each pass, your worst case is O(n^2). 

The **best case** is `O(n)` if the list is **already sorted** as then no swaps are needed, and only one pass of the inner loop takes place and due to no swaps triggering the break and return. 

The **average case** I found harder to work out as initially I thought that the break optimization may help make this better than `O(n^2)` but that break only triggers when the list is already sorted on a pass. How likely is that for an average random list? I first thought about a smaller example case for a random list of `5 elements` [32, 64, 16, 45, 21] and worked out the distances that each element of the list had from its final position and got that the average distance = `(2 + 3 + 2 + 0 + 3)/5` = 10/5 = `n/2`. I then **generalized this** since **each item** of `n` could be **equally likely** to be anywhere, the **average distances** = `(0 + 1 + 2 + 3.....(n-1))/n` and summing that I used the trick from **Gauss/Triangular numbers** that becomes `n(n-1)/2/n` = `(n-1)/2` ~ `n/2`. Then if each element is n/2 places away from its final position and each element can move one step per pass then we would need n/2 passes before elements are able to trigger the break so that is `~ O(n)` passes since with Big O we can drop constants as this shows it still scales with n. Each pass also has an inner loop where we range from `j = 0` to `j = total - i-1` where `total` is `n` so the number of passes shrinks as we go but is still **proportional** to `n` as we have `n` within the **upper bound** so then comparisons per pass would average to  ~ `n/2` so `~ O(n)`. If we then combine this estimate with the outer loop it comes to: `O(n) x O(n)` (dropping the constants /2) so still `O(n^2)` for the average list case. 


#### Space complexity
The space complexity for the function I defined is `O(n)` since as per the exercise brief I had to create a new list each time that `bubble_sort` is called. If I could have sorted the list in place space complexity could have been `O(1)`. 

#### Stability 
The bubble sort function I defined is stable since I only used greater than in the inner loop and not greater than or equal to, so a swap only happens if the element at `j` is larger than the element at `j+1`. If the elements at `j` and `j+1` are `equal`, then their `original order is preserved`, resulting in a stable operation.  

##### Reference:
- Triangular number:  https://en.wikipedia.org/wiki/Triangular_number

### Merge Sort: - ASH TO FIX LAST TWO 

#### Time complexity
For the **worst case** where all elements need to be reviewed in `merge_sort` we split the list into left and right and then each split list left and right keep being split and the **number of levels** this goes to is from `n to 1` so `log₂(n)`. Then for the **merge function** we pull those left and right lists back together again placing each element one at a time from one of the lists and then appending the remainder, the leftover item, from the other, so that would scale with n so that is `O(n)`. If we combine this then each level has `O(n)` and we have `O(log₂n)` levels due to recursion. This simplifies to `O(nlogn)` since changing bases of logs is just constant multiplication and for Big O we know we can ignore constants . 

What about the **best case**, an already sorted list. This will still have left and right splits run through the recursive loops as the recursion only stops when the base case is hit when length is 1 so even if the list is sorted you still have `O(logn)` for the levels of your recursion and then the **merge** will be the same as the worst case as even if our list is sorted the `merge` function will still need to place each element from one of the lists once (and then append the remainder from the other). I initially fell into the trap of thinking oh this is faster since there are fewer comparisons which there are in the best case but Big O complexity is about scaling with n and the append is still a step that will add more items as n grows so like the worst case, will scale with n so best case is then also `O(nlogn)`.

Since we established that the merge function will always need to do n placements regardless of how sorted the input is, there is no best case optimization and so the  **average case** would also be `O(nlogn)` time complexity. 

#### Space complexity
Looking at my algorithm we first have our merge_sort function, here we have our two half lists O(n/2) each so O(n) and the recursive function calls so we have a **stack** and the deepest it can go (halving each time) is from `n to 1` so `log₂(n)` which is ~ `O(logn)`. We then combine our lists to form the merged list, `total of size n`. I initially thought this would be then O(nlogn) but then I realized each level does not hold a list of n elements, the deeper the level of recursion the fewer elements of n there are which simply together would sum up to ~n rather than each holding n elements. So we have `O(logn)` for the recursive stack and O(n) for the half lists and final list and `O(n)` is larger so that would be our peak amount of memory we have to hold at a point so space complexity is `O(n)`. 

**Additional research, beyond scope of the task:** The exercise asked to create a new list and **not to sort in place** but I was **curious** if it was possible to do a `merge sort` in place and see if that would improve **space complexity** to O(1). Through research I found out about `block merge sort` which is a type of merge sort which is **non-recursive** so there is no stack and which simplifies to `O(1)` space complexity. 

##### Reference: 
- Block merge sort: https://en.wikipedia.org/wiki/Block_sort

#### Stability
In the input unsorted list the elements in the left half would have come before the elements in the right half. I added the condition if `if left[i] <= right[j]:` **append from left list first**, so then if any of the items in each half are equal (the `=` part of `<=`),  the **left items** will be added first so the order will remain which guarantees stability so my `merge sort` implementation ** is stable**. 

<!-- grade_id: cell-66f388f6c90d3698, points: 0, grade: false, solution: true -->
```python
def prefix_to_postfix(expression):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0a52d49715d5f7f1, points: 25, grade: true, solution: false -->
```python
assert(prefix_to_postfix("5") == "5")
assert(prefix_to_postfix("+ 10 20") == "10 20 +")
assert(prefix_to_postfix("* + 1 2 - 5 3") == "1 2 + 5 3 - *")
```